# Worked Example 04

In this exercise, you will practice wrangling data through different formats, with missing data, and working with text.

## Autograder Instructions ##

Run the following cell to install the Otter autograder package the first time it is run. Afterwards, running it again will print out its version number if Otter is correctly installed.

```python
try:
    import otter
    grader = otter.Notebook("worked-example04.ipynb")
    print(otter.__version__)
except Exception as e:
    %pip install otter-grader
```

Output:
```
6.1.6
```

```python
import numpy as np
```

To grade your own work, simply run the cell starting with grader.check immediately after each question. The cell calls Otter to run the tests for all subquestions and generate a report of what test(s) you pass/do not pass in the same style as on Gradescope.

## Part 1: Text wrangling and regular expressions 

### Question 1
Below we import the `text.txt` file as a string called `text` and preview the first 1,000 characters. Note that the text has no capitalization at the beginning of sentences.

```python
# Run but do not modify this code
with open("text.txt", encoding="utf-8") as f:
    text = f.read()

print(text[:1000])
```

Output:
```
Cell In[3], line 2
      1 # Run but do not modify this code
----> 2 with open("text.txt", encoding="utf-8") as f:
      3     text = f.read()
      5 print(text[:1000])

File ~/Library/Python/3.12/lib/python/site-packages/IPython/core/interactiveshell.py:344, in _modified_open(file, *args, **kwargs)
    337 if file in {0, 1, 2}:
    338     raise ValueError(
    339         f"IPython won't let you open fd={file} by default "
    340         "as it is likely to crash IPython. If you know what you are doing, "
    341         "you can use builtins' open."
    342     )
--> 344 return io_open(file, *args, **kwargs)

FileNotFoundError: [Errno 2] No such file or directory: 'text.txt'
```

Capitalize the first letter of each sentence (e.g., "it was the best..." should become "It was the best..."). To do so, split the sentences (you may assume sentences are ended with the a period `.` followed by a blank space), capitalize the first letter of each sentence, and then join the results.

Put your result in the variable `q1`.


```python
# Put your code to answer the question here
# Please store your answer as the givin variable

q1 = ...


# Run but do not modify this code
start = q1.find('.')
print('q1 section:', q1[start-5:start+10])
```

```python
grader.check("q1")
```

### Question 2
Below we import the `corrupted_text.txt` file, a messed up version of `text.txt`, and print the first 1,000 characters.

```python
# Run but do not modify this code
with open("corrupted_text.txt") as f:
    ctext = f.read()

print(ctext[:1000])
```

Fix the following problems to recover the original text that appeared in `text.txt`.

1. Every comma `,` has been replaced by `$`. Sometimes, but not always, this is followed by `####`. When it is followed by `####`, there are different numbers of blank spaces after the `####`.
2. Every period `.` has been replaced by `$\n\n`.

When you are finished, if you print the first 1,000 characters, you should see the same text as in `text.txt` above (except without the capitalization you worked on in the previous problem).

Put your result in the variable `q2`.

Note: It may be helpful to recall that `\` is an escape character for Python strings, so for example, while calling `print("\n")` will simply print a blank newline character, calling `print("\\n")` will "escape" the usual meaning of the second backslash and print `\n`.


```python
# Put your code to answer the question here
# Please store your answer as the givin variable

q2 = ...



# Run but do not modify this code for debugging help
print('q2 section:', q2[:100])
```

```python
grader.check("q2")
```

### Question 3
Below, we again import the `text.txt` file as a new string variable `ftext` (for "fresh" text if that helps).

```python
# Run but do not modify this code
import re # Regular expression library
with open("text.txt") as f:
    ftext = f.read()
```

Output:
```
Cell In[6], line 3
      1 # Run but do not modify this code
      2 import re # Regular expression library
----> 3 with open("text.txt") as f:
      4     ftext = f.read()

File ~/Library/Python/3.12/lib/python/site-packages/IPython/core/interactiveshell.py:344, in _modified_open(file, *args, **kwargs)
    337 if file in {0, 1, 2}:
    338     raise ValueError(
    339         f"IPython won't let you open fd={file} by default "
    340         "as it is likely to crash IPython. If you know what you are doing, "
    341         "you can use builtins' open."
    342     )
--> 344 return io_open(file, *args, **kwargs)

FileNotFoundError: [Errno 2] No such file or directory: 'text.txt'
```

Write regular expressions and use the Python `re` module [documentation](https://docs.python.org/3/library/re.html) to answer the following questions about the file. You may assume that all characters are lowercase. You can print your answers directly from code (as long as they are clearly labeled) or you can write them in the markdown cell below.

1. How many times do two vowels (a, e, i, o, and u) appear consecutively in `text`? When three vowels occur consecutively (such as in "queen") you only need to count once, (for example, in "queen" you do not need to worry about the overlap between "ue" and "ee", just count once). 
2. Find all of the hyphenated words in `text` (that is, the words with a hyphen character `-` in the middle and no spaces between the hyphen and the characters on either side). 
3. Find the first word that appears after each period character `.` in the text. Do not include the period itself or the space that follows it. They should all be lower case from the `ftext`.
4. Find the starting index of every word `she` or `he`. Do not include words that enclose these such as `the` or `fishes` (Hint: It appears only 3 times).

Put question 1's answer in the variable `q3_1`, question 2's answer in `q3_2`, and so on.


```python
# Put your code to answer the question here
# Please store the answer for each questions as following variables

q3_1 = count 
q3_2 = ... 
q3_3 = ...
q3_4 = ...




# Run but do not modify this code for debugging help
print('q3_1:', q3_1)
print('q3_2:', q3_2)
print('q3_3:', q3_3)
print('q3_4:', q3_4)
```

```python
grader.check("q3")
```

## Part 2: Cleaning up system logs CSV
In this part we work with a piece of messy tabular data in the form of a poorly formatted csv file containing data about programs running on computer systems. 

### Question 4
Below, we import the dataset using the Pandas `read_csv` function that creates a dataframe. Run the code; it will preview the first five rows. There are several formatting issues with the default import.

```python
# Run but do not modify this code
import pandas as pd
sys_df = pd.read_csv("monitor.csv")
sys_df.head()
```

Write a new call to `read_csv` that addresses the following.

1. The rows are already numbered with an integer index, but that is being duplicated so that we see a primary index for the dataframe and separately a column containing integer indices. Fix this so that there is just a single integer index as the primary index of the dataframe (bolded). Hint: You may find calling `read_csv` with specific arguments easier than cleaning `sys_df` created in the prior cell.

2. There are commas at the end of every row causing the dataframe to add a blank extra column at the right. Get rid of this column.

3. The file did not have column labels, so currently the first row is being incorrectly interpreted as the bolded column labels. Fix the code so that the column labels (from left to right) are `System Time second`, `System Memory GB`, and `System Memory MB`.


When you are finished, `sys_df` should have the above issues correct. Assign the variable `q4` with the same value as `sys_df`.


```python
# Put your code to answer the question here


# Run but do not modify this code
q4 = sys_df
q4.head(10)
```

```python
grader.check("q4")
```

### Question 5
The `sys_df` dataframe from question 3 should now be a little easier to read and use, but there are some strange things we would like to fix about the formatting of the data. When you are finished, `sys_df` assign the variable `q5_1` with the same value as `sys_df`.

**Do not** use for loops, unless it is guaranteed to iterate less than 50 times. In the homework there will be manual points that check for this. You are also likely to not pass some tests because iterating over the entire DataFrame sometimes causes issues with formatting.

1. The three features or columns all contain numeric data (seconds, GB, and MB) but are currently formatted as Strings with redundant information also contained in the column names. Fix the dataframe so that each value is just an integer. Note that the rows at index `400` and on have System Time recorded in minutes instead of seconds, be sure to convert these to seconds by multiplying by 60.

2. Currently the System Memory is split accross two columns, one for the GB and one for the MB. For example, the total memory of the first program is 3 GB and 414 MB. Instead, represent the full system memory in the `System Memory MB` column, and get rid of the `System Memory GB` column. To do so, you need to convert the values in the GB column to MB (1 GB is 1000 MB) and add that to the MB column, then use the [`drop` method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.drop.html) to remove the GB column.

3. Now that you have dealt with the above, it should be easy to provide some basic statistics about the data. Report the mean time in seconds in the variable `q5_2` and mean memory in MB of the programs in the variable `q5_3`.


```python
# Put your code to answer the question here

q5_1 = ... # q5_1 = sys_df, after the manipulation
q5_2 = ... # the mean time in seconds
q5_3 = ... # the mean memory in MB



# Run but do not modify this code
q5_1 = sys_df

# Run but do not modify this code for debugging help
print('q5_2:', q5_2)
print('q5_3:', q5_3)
q5_1.head(10)
```

```python
grader.check("q5")
```

## Part 3: Wrangling FDA JSON Dataset 
In this part we work with a messy JSON dataset containing information about several drugs labels.

### Question 6
Below we import the `FDADrugLabel.json` file into the `labels` variable.

The resulting Python object is somewhat messy; we encourage you to explore the data before answering the questions. For example, rather than just printing out `labels` check it's type, print only one element, and remember there are different [dictionary functions](https://docs.python.org/3/library/stdtypes.html#typesmapping) you could use. Feel free to add your own cells as you explore.

```python
# Run but do not modify this code
import json
with open("FDADrugLabel.json") as f:
    labels = json.load(f)
```

Answer the following questions. You may use `for` loops here since this is outside of the `pandas` and `numpy` libraries. Put question 1's answer in the variable `q6_1`, question 2's answer in `q6_2`, and so on.

1. How many drugs are represented in the file?
2. What is the largest number of key/value (or name/value) pairs for any drug in the file?
3. The `effective_time` is recorded for each drug as a date in the format `yyyymmdd`, for example, June 1, 1992 would be recorded as `19920601`. What is the earliest recorded date?


```python
# Put your code to answer the question here

q6_1 = ... # number of drugs
q6_2 = ... # largest number of pairs
q6_3 = ... # earliest recorded date

...

# Run but do not modify this code for debugging help
print('q6_1:', q6_1)
print('q6_2:', q6_2)
print('q6_3:', q6_3)
```

```python
grader.check("q6")
```

### Question 7

Below we use the Pandas `read_json` function to import the `FDADrugLabel.json` file into a Pandas dataframe `label_df`.

```python
# Run but do not modify this code
import pandas as pd
label_df = pd.read_json("FDADrugLabel.json")
label_df.head()
```

Use the dataframe to answer the following questions.

1. What happens if you call `label_df.dropna()` (with no optional parameters) to remove the null values? Why? 
 
2. In `q7_2` put all of `label_df`'s columns that contain data for every row. You might answer this using the `label_df.dropna()` method to keep all of the rows but remove any of the columns that have any missing data, but do not modify the `label_df` dataframe itself. You may use `list` or a `for` loop to create this list

3. Filter the dataframe to the rows that are missing data for `active_ingredient`, then put them in `q7_3` (if `my_view` is a dataframe view you have computed by filtering the data, you can simply put `my_view` as the last line of a cell and run it to get a tabular rendering of the rows in Jupyter). You may find the Pandas `isnull` method useful, you can see it's [documentation here](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.isnull.html).

<!-- BEGIN QUESTION -->

### Put your answer to question 1 here:

<!-- END QUESTION -->



```python
# Put your code to answer the question here
# feel free to add more cells of code is needed

q7_2 = ... # columns cotains data for every row. Please store the answer as a list of column names
q7_3 = ... 




# Run but do not modify this code for debugging help
print('q7_2:', q7_2)
print('q7_3:', len(q7_3))
```

```python
grader.check("q7")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: [1, 1])
```
>>> '. i' not in q1
True
```
```
>>> bool(q1[:1000] == 'It was the best of times, it was the worst of times, it was the age of wisdom, it was the age of foolishness, it was the epoch of belief, it was the epoch of incredulity, it was the season of light, it was the season of darkness, it was the spring of hope, it was the winter of despair, we had everything before us, we had nothing before us, we were all going direct to heaven, we were all going direct the other way- in short, the period was so far like the present period, that some of its noisiest authorities insisted on its being received, for good or for evil, in the superlative degree of comparison only. There were a king with a large jaw and a queen with a plain face, on the throne of england; there were a king with a large jaw and a queen with a fair face, on the throne of france. In both countries it was clearer than crystal to the lords of the state preserves of loaves and fishes, that things in general were settled for ever. It was the year of our lord one thousand seven hundred ')
True
```


Tests for q2 (points: [0.4, 0.4, 0.4, 0.4, 0.4])
```
>>> '$' not in q2
True
```
```
>>> '####' not in q2
True
```
```
>>> '\n\n' not in q2
True
```
```
>>> '  ' not in q2
True
```
```
>>> bool(q2[:1000] == 'it was the best of times, it was the worst of times, it was the age of wisdom, it was the age of foolishness, it was the epoch of belief, it was the epoch of incredulity, it was the season of light, it was the season of darkness, it was the spring of hope, it was the winter of despair, we had everything before us, we had nothing before us, we were all going direct to heaven, we were all going direct the other way- in short, the period was so far like the present period, that some of its noisiest authorities insisted on its being received, for good or for evil, in the superlative degree of comparison only. there were a king with a large jaw and a queen with a plain face, on the throne of england; there were a king with a large jaw and a queen with a fair face, on the throne of france. in both countries it was clearer than crystal to the lords of the state preserves of loaves and fishes, that things in general were settled for ever. it was the year of our lord one thousand seven hundred ')
True
```


Tests for q3 (points: [2, 2, 2, 2])
```
>>> bool(q3_1 == 184)
True
```
```
>>> print(sorted(q3_2))
['cock-lane', 'cock-lane', 'creatures-the', 'drawing-rooms', 'fellow-tradesman', 'five-and', 'rest-along', 'seventy-five', 'seventy-five', 'seventy-five', 'to-day', 'to-morrow']
```
```
>>> print(sorted(q3_3))
['all', 'but', 'daring', 'environed', 'even', 'france', 'giles', 'in', 'in', 'in', 'it', 'it', 'it', 'mere', 'mrs', 'southcott', 'spiritual', 'there', 'thus', 'under']
```
```
>>> bool(set(q3_4) == {2087, 2269, 3760})
True
```


Tests for q4 (points: [1, 1, 1, 1, 2])
```
>>> '0' not in q4.columns
True
```
```
>>> ' ' not in q4.columns
True
```
```
>>> bool(len(q4.columns) == 3)
True
```
```
>>> bool(len(q4) == 1000)
True
```
```
>>> bool(set(q4.columns) == set(['System Time second', 'System Memory GB', 'System Memory MB']))
True
```


Tests for q5 (points: [1, 1, 1, 1, 0.5, 0.5, 0.5, 0.5])
```
>>> bool(len(q5_1.columns) == 2)
True
```
```
>>> bool(q5_1['System Time second'].dtypes == 'int64' and q5_1['System Memory MB'].dtypes == 'int64')
True
```
```
>>> bool(set(q5_1.columns) == {'System Memory MB', 'System Time second'})
True
```
```
>>> bool(list(q5_1['System Time second'][[0, 200, 450, 600]]) == [23, 38, 2220, 1800] and list(q5_1['System Memory MB'][[0, 200, 450, 600]]) == [3414, 2415, 2483, 2404])
True
```
```
>>> isinstance(q5_2, float)
True
```
```
>>> bool(round(q5_2, 5) == 1172.581)
True
```
```
>>> isinstance(q5_3, float)
True
```
```
>>> bool(round(q5_3, 5) == 2869.839)
True
```


Tests for q6 (points: [1, 1, 1, 1, 1, 1])
```
>>> isinstance(q6_1, int) or isinstance(q6_1, np.int64) or isinstance(q6_1, np.int32)
True
```
```
>>> bool(q6_1 == 20)
True
```
```
>>> isinstance(q6_2, int)
True
```
```
>>> bool(q6_2 == 36)
True
```
```
>>> isinstance(q6_3, str)
True
```
```
>>> bool(q6_3 == '20150707')
True
```


Tests for q7 (points: [1, 1, 1, 1])
```
>>> isinstance(q7_2, list)
True
```
```
>>> bool(set(q7_2) == set(['package_label_principal_display_panel', 'effective_time', 'openfda', 'spl_product_data_elements', 'set_id', 'version', 'id', 'indications_and_usage']))
True
```
```
>>> bool(len(q7_3) == 3)
True
```
```
>>> bool(set(q7_3['set_id']) == {'4146cfc6-2b02-4731-a0f1-8c445bafea55', 'b84ccb27-9c52-43b9-9a46-8f52a571045e', 'c3e0a3b9-535d-456b-9a1f-fa18a4d4e38f'})
True
```
